```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment_06.ipynb")
```

# File Names, Project Organization and Virtual Environments


<!-- BEGIN QUESTION -->

<div class="alert alert-info">

## Submission instructions

<p>You receive marks for submitting your lab correctly, please follow these instructions:</p>

<ul>
  <!-- <li>Make at least three commits.</li>-->
  <li>Export to HTML via <code>File -> Save and export notebook as... -> HTML</code> and upload the HTML file to Canvas.</li>
    <ul>
      <li>Before submitting, make sure you restart the kernel and rerun all cells.</li>
    </ul>
  <li>Don't change any variable names that are given to you, don't move cells around, and don't include any code to install packages in the notebook.</li>
  </li>
</ul>
</div>

_Points:_ 1

<!-- END QUESTION -->

----

## Recap

**Step 1** Let's first practice skills we have learnt in a previous Modules:

1.  Create a new GitHub repo called `project`

    - make repo `Public`
    - initialize a repo with `README` file
    - Choose a license : `MIT License`  

2.  After the repo is created, add `.gitignore` file and populate it with `**/*.log`, `.DS_Store` and `.ipynb_checkpoints`
 
 
3. Clone the `project` repo to your computer

    - create a branch called `developer-1` 
    - set up the tracking branch for your local `developer-1` branch and use it to make any changes to your repo 
    - publish the branch on the remote



4. Fork the repo https://github.com/yuliaUU/files_to_copy 

    - copy the **content** of the repo into your `project` repo (**do not add any hidden or git files**)


   
for this exercise you will be working with your `project` repo and files that you have copied

-----

## Part I: File Naming

A proper file naming system is essential when working on a project to ensure organization, efficiency, and collaboration. Consistent naming conventions make it easier to locate and identify files, reducing time spent searching and minimizing the risk of losing important data. Additionally, a well-structured naming system promotes seamless collaboration among team members, as everyone can easily understand and access the required files, thereby streamlining the overall project workflow.

<div class="alert alert-info">
    
### Question 1

Which of the following are important rules for naming files on a computer? Select all correct answers.


A)  Files names should play well with default ordering  
B)  The American date convention MM-DD-YYYY is good to use in file names  
C)  File name parts should be easy to extract programmatically  
D)  Underscores are the only allowed separator in file names  
E)  File names should always use upper-case letters  
F)  File names should be informative to humans  
</div>
 

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1 = "ACF"
q1
```

Output:
```
'ACF'
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 2
    
Match each explanation with its corresponding file naming principle discussed in this module:
    
<table>
<tr>
<td> <b>Definitions:</b></td>
<td> <b>Match the terms on the left with the definitions from the list below:</b></td>
</tr>
<tr>
<td> 
    
1. Machine readability
2. Human readability
3. Compatibility with default ordering

   
</td>
<td> 
        
A)  By following this principle, we will be sure that our files have the appropriate formatting in terms of their content. Such files will be ready for subsequent analysis in a data science pipeline.  
B) By following this principle, we can have a good idea of what is inside files without having to open them (saves us time) and when we come back to work on a project we haven't worked on in a while it is much easier for us to remember what we were working on and get started again.    
C) By following this principle, we can make our files more organized. Files will be sorted into a logical order so it is easy to orient yourself in the directory and find the files you want to find.   
D) By following this principle, we are able to use tools such as regular expressions and file-name globbing to search, select and manipulate files using tools such as the unix shell, R and Python.     

</td>
</tr>
</table>
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q2 = "DBC"
q2
```

Output:
```
'DBC'
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 3
    
Open the `project` repo and navigate to `weekly meetings` folder. There you can see a list of five `.md` files with meeting minutes for a data science company. Which of the following file names are best in line with the principles you've learned?

A) `meeting-minutes_2020-01-21T0700.md`    
B) `2023-1-21_meeting-minutes.md`    
C) `2023-01-21T0700_meeting-minutes.md`    
D) `2023-01-21T07:00_meeting-minutes.md`  
E) `2023-01-21T0700meetingminutes.md`  
 

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q3 = "C"
q3
```

Output:
```
'C'
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Step 2:** Rename the files in your `weekly meetings` repository according to the best practices for file names. Commit your changes and push to the remote

<div class="alert alert-info">
    
### Question 4 

Go to folder `Documents`. Which of the following file names are best in line with the principles you've learned?

A) `Report 01.docx`   
B) `Report1.docx`  
C) `Report01.docx`   
D) `report_01.docx`     
E) `report 01.docx`   

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q4 = "D"
q4
```

Output:
```
'D'
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

**Step 3:** Rename all files in `Documents` folder accordingly, commit and push changes

<div class="alert alert-info">
    
### Question 5

Find `dataset-csv` file  in the repository. What issues can you find with it?



A)	There are no issues, the name is comprehensive and clear  
B) The file extension format is off  
C)	The file name is not machine readable   
D)	The file name is not informative to humans  
E)	The file name should start with an upper-case letter  


</div>


_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q5 = "B"
q5
```

Output:
```
'B'
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

**Step 4:** Rename the file accordingly, commit and push changes

<div class="alert alert-info">
    
### Question 6

Find all python scripts in the folder (`.py` extention). These files are run in the following order: generate the data with `gen-data.py`, plot the dataset `untitled1.py` and plot the predictions with `untitled2.py`
    
What will be a better way to name the files?
    
A)	`generate-data_01.py`, `visualize-data_02.py`, `plot-predictions_03.py`  
B)  `01generate-data.py`, `02visualize-data.py`, `03plot-predictions.py`   
C)	`script_01.py`, `script_02.py`, `script_03.py`  
D)	`01_generate-data.py`, `02_visualize-data.py`, `03_plot-predictions.py`  
  
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q6 = "D"
q6
```

Output:
```
'D'
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Step 5:** Rename these files accordingly in your repo, commit and push changes

-----

# Part II: Project organization

Once you done with renaming the files, let's make sure that our project organization is clean and easy to follow. 

One very useful learning strategy for data science is to look and reflect on how others in the field do things, and then incorporate the useful ideas/strategies/tools into your own practice. Let's look at project organization for two different types of data science projects:

-   data analysis projects -> the focus is answering a research/business question using data science tools,
-   data science tools projects -> the focus is creating a tool you or others can later use to help solve a research/business question with data analysis.

Here are some example repositories for both categories of projects:

Data analysis:

-   <https://github.com/benmarwick/Pleistocene-aged-stone-artefacts-from-Jerimalai--East-Timor>
-   <https://github.com/bici-sancta/gdp_worldbank>
-   <https://github.com/jooolia/RdRp_454_amplicons_Jericho_and_SOG>
-   <https://github.com/ttimbers/equine_numbers_value_canada>

Data science tools:

-   [dplyr](https://github.com/tidyverse/dplyr)
-   [seaborn](https://github.com/mwaskom/seaborn)
-   [pandas](https://github.com/pandas-dev/pandas)
-   [topojson](https://github.com/topojson/topojson)

<div class="alert alert-info">
    
### Question 7

Now briefly reflect on the similarities and differences between the project organization used in the analysis and tools repos. In doing so, consider both the folder organization and which files go into the folders. For each statement (1-6) match it with a letter (A, B, or C) to indicate which features are more common in which type of project.

<table>
<tr>
<td> <b>Definitions:</b></td>
<td> <b>Match the terms on the left with the definitions from the list below:</b></td>
</tr>
<tr>
<td> 
    
1. Having a license file and contribution guidelines
2. Containing a report, presentation, or manuscript
3. Using folders to organize the projects rather than a flat hierarchy
4. Having documentation and setup/installation instructions
5. Having a data folder
6. Having some sort of environment file listing dependencies for increased reproducibility
    
</td>
<td> 
        
A) More common in data analysis projects  
B) More common in data science tools projects     
C) Common in both types of projects  

</td>
</tr>
</table>

</div>

_Points:_ 1

```python
# Your answer should be a string with 6 letters of the options you think are correct, e.g. 'ABACCA'
q7 = "BACBAC"
q7
```

Output:
```
'BACBAC'
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 8 

Which repo structure you think fits our `project` the best?

<table>
<tr>
<td>

A) 
```
project/
├── data/
│   ├── dataset-csv
│   ├── dataset-image.png
│   └── LICENSE.md
├── docs/
│   ├── Documents/
│   ├── Report.pdf
│   └── Report.ipynb
├── predictions/
│   └── predictions.png
├── src/
│   ├── gen-data.py
│   ├── untitled1.py
│   └── untitled2.py
└── weekly-meetings/
└── README.md
```

</td>
<td>

B)
```
project/
├── data/
│   ├── dataset-csv
│   └── dataset-image.png
├── docs/
│   ├── Documents/
│   ├── Report.pdf
│   └── Report.ipynb
├── predictions/
│   └── predictions.png
├── src/
│   ├── gen-data.py
│   ├── untitled1.py
│   └── untitled2.py
└── weekly-meetings/
├── LICENSE.md
└── README.md
```
    
</td>
</tr>
<tr>
<td>
  
C)
```
project/
├── data/
│   ├── dataset-csv
│   ├── dataset-image.png
│   └── docs/
│       ├── Documents/
│       ├── Report.pdf
│       └── Report.ipynb
├── src/
│   ├── gen-data.py
│   ├── untitled1.py
│   └── untitled2.py
├── weekly-meetings/
└── predictions.png
├── LICENSE.md
└── README.md
```

<br> 
<br>    
</td>
<td>

D)
```
project/
├── data/
│   └── dataset-csv
├── docs/
│   ├── Documents/
│   ├── weekly-meetings/
│   └── Report.pdf
├── images/
│   ├── dataset-image.png
│   └── predictions.png
├── src/
│   ├── gen-data.py
│   ├── untitled1.py
│   └── untitled2.py
├── reports/
│   └── Report.ipynb
├── LICENSE.md
└── README.md
```

</td>   
</tr>
</table>
      

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q8 = "D"
q8
```

Output:
```
'D'
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

**Step 6:** Organize your files accordingly, make sure that all files/ folders are meeting requirements of proper naming and organization, commit the changes and push your `project` to remote

-----

# Part III: Conda  

Let's explore the scenario: Imagine you have a desk want to keep different species of snakes. Now, each species of snake requires different conditions to survive, such as temperature, humidity, and lighting. Thus, you need multiple terrariums with different species of snakes in it. You can customize each terrarium with its own set of conditions to suit the species of snake. You also need a pet store where you can buy the required food and accessories for your snakes. 

<img src="images/Picture1.png" width=80% >

In this scenario:
 - Your desk represents your computer
 - Each terrarium with snakes represents a project or a set of applications that you are working on. Similarly, each project or set of applications that you work on might require different versions of programming languages (i,.e different versions of Python), libraries, and dependencies. This is where Conda virtual environments come in.
 - like a terrarium, A **conda virtual environment** is like a separate container that isolates the dependencies and libraries required for a particular project or application from the rest of the system. It allows you to create and manage multiple environments, each with its own set of dependencies, without interfering with other projects or applications. Each virtual environment is like a separate terrarium with its own set of conditions that can be customized to suit the species of snake.
- Next, we have **channels**. Channels in Conda are like pet stores where you can browse and download packages that you need for your projects. There are public channels available, like anaconda and conda-forge, and private channels that you can create for your own organization.
- Finally, we have **packages**. Packages are like the food and accessories that you buy for your snakes. They are the collections of libraries and dependencies that you need for your projects. You can download packages from channels and install them in your virtual environments.



<div class="alert alert-info">
    
### Question 9

Which of the following descriptions is most accurate about Conda?

A)  Conda is a package and environment manager for many languages   
B)  Conda is an environment manager only for Python packages used in Jupyter Lab   
C)  Conda is a package manager for all languages   
D)  Conda is a tool similar to Git for version-controlling Python packages   
E)  Conda is a package and environment manager only for Python   
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q9 = "A"
q9
```

Output:
```
'A'
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 10

Which of the following statements are correct about Conda?

A)  It only works on Linux and Mac at the moment     
B)  It is a command-line tool     
C)  It it installed as part of Miniconda and Anaconda    
D)  It ensures the reproducibility of a computational environment    
E)  It requires JupyterLab to access different environments   
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q10 = "BCD"
q10
```

Output:
```
'BCD'
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

## Conda Environments


<div class="alert alert-info">
    
### Question 11


Conda environments help you isolate packages that you use for specific projects and automatically handle dependencies and version requirements for each set of packages.

A)  True   
B)  False
  

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q11 = "A"
q11
```

Output:
```
'A'
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 12

One downside of Conda environments is that the same version of Python is shared among all environments.

A)  True   
B)  False

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q12 = "B"
q12
```

Output:
```
'B'
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

**Step 7:** Let's now practice creating a environment for our `project`. Create the environment called `dstoolbox`

<details><summary><b>Hint:</b></summary>
You will need to use <code>conda create</code> command
</details>

Double check that you created the environment successfully by running `conda info --envs`. You should see `dstoolbox` in a list of environments (most likely you just have two at this point) 

<div class="alert alert-info">
    
### Question 13
    
We would like to use the created `dstoolbox` environment in JupyterLab. In which of the following scenarios will we be able to use the `dstoolbox` environment in JupyterLab? Select all correct answers.

A)  Install `jupyterlab` in the base environment  
B)  Install `jupyterlab` in the `dstoolbox` environment  
C)  Install `jupyterlab` and `ipykernel` in the base environment, and `nb_conda_kernels` in the ml-model environment.   
D)  Install `jupyterlab` and `nb_conda_kernels` in the base environment, and `ipykernel` in the `dstoolbox` environment  
E)  As long as JupyterLab is installed in one of the existing Conda environments, it will automatically have access to all other environments on the same system. 
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q13 = "BD"
q13
```

Output:
```
'BD'
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
```

**Step 8:**  Follow the steps described in Module 6 to make Conda available in Jupyter Lab

<details><summary><b>Hint:</b></summary>

To make your new environment visible in Jupyter Lab you need the following steps: 
    
1. install <code>nb_conda_kernels</code> into your **base** environment
2. install <code>ipykernel</code> into your **dstoolbox** environment
    
</details>
  

After you completed this step, open Jupyter Lab to check if your new environment is available using `conda env list` command. Activate your new `dstoolbox`environment.

----

## Channels & Packages

Conda **channels** are the locations where packages are stored. A Conda channel is essentially a collection of Conda packages hosted on a remote server, which can be accessed and installed using the conda package manager.

The `default` Conda channel is hosted on the Anaconda Cloud platform, which contains a large collection of commonly used packages for scientific computing, data analysis, and machine learning. *Often, we only need to use the default channel* as it contains so many of the commonly used data science packages. However, if a package is not available on the default channel, you can search for it on other channels or create your own channel.

There are several public Conda channels available, such as:

- `conda-forge` which contains a large collection of community-contributed packages, including more recent versions of many of the packages in the default channel
- `bioconda` which contains packages for bioinformatics and computational biology
- `pytorch` which contain many machine learning packages

To use a package from a Conda channel, you can specify the channel name when installing the package using the `conda` command. For example, to install a package from a specific channel, you would run `conda install -c channel_name package_name`

Channels provide access to different versions of **packages**, ensuring compatibility and enabling users to choose the most appropriate version for their needs. Different versions of a package refer to distinct releases of a software library or application that include updates, improvements, bug fixes, or new features compared to previous iterations. As developers work on a package, they typically assign version numbers to indicate the progress and evolution of the software, helping users identify the most suitable version for their requirements and ensuring compatibility with other packages or system components.

For example, package versioning:

- 2.0.0: This is a major release where new features or significant changes are introduced, potentially causing incompatibilities with previous versions (1.x.x).
- 2.1.0: This would be a minor release, adding new features or improvements while maintaining backward compatibility with the 2.x.x series.
- 2.1.1: This represents a patch release, fixing bugs or making minor updates without adding new features, and maintaining compatibility with version 2.1.x.

In the case of pre-release versions:

- 3.0.0-alpha: This is an alpha release for version 3.0.0, indicating that it's an unstable version under development and not recommended for production use.
- 3.0.0-beta: This is a beta release for version 3.0.0, suggesting that the software is closer to being stable but may still have some bugs or issues.


<div class="alert alert-info">
    
### Question 14

We would like to install a package called `scikit-learn` in our environment, but we first want to check what versions are available on different channels. Your task is to use the appropriate Conda command to search for `scikit-learn` in `pkgs/main` channel.

Write a shell command that produces a one-line output showing the latest version of `scikit-learn` in the `pkgs/main` channel. Here's an example output:

`scikit-learn  1.2.2  py39hd77b12b_0  pkgs/main`


<details><summary><b>Hint:</b></summary>
Consider using shell command such as <code>grep</code>, and the pipe operator that you've learned about in previous modules to achieve this. To see the last line you can use <code>tail -n 1</code>
</details> 

</div>

_Points:_ 1

```python
# Your answer should be a string with a single line command surrounded by ''
q14 = ...
```

```python
grader.check("q14")
```

You can install a package into an existing environment using the `conda install` command. This command accepts a list of package specifications (i.e., `scikit-learn=0.22`) and installs a set of packages consistent with those specifications and compatible with the underlying environment. If full compatibility cannot be assured, an error is reported and the environment is not changed.

By default the conda install command will install packages into the current, active environment.
As was the case when listing packages to install when using the `conda create` command, if version numbers are not explicitly provided, Conda will attempt to install the newest versions of any requested packages. To accomplish this, Conda may need to update some packages that are already installed or install additional packages. It is always a good idea to explicitly provide version numbers when installing packages with the conda install command. For example, the following would install a particular version of scikit-Learn, into the current, active environment.

`conda install scikit-learn=0.22`


<div class="alert alert-info">
    
### Question 15
    
What command we can use to install `scikit-learn` 1.2.1 from the `pkgs/main` channel into the currently activate environment?
    
Use the shortest possible command that specify both package version and channel
    
</div>

_Points:_ 1

```python
# Your answer should be a string with a single line command surrounded by ''
q15 = ...
```

```python
grader.check("q15")
```

**Step 9:** Install version 1.2.1 of`scikit-learn` package from the `pkgs/main` channel into the currently activate environment 

## Exporting Conda Environment

<div class="alert alert-info">
    
### Question 16
    
Which of the statements below are correct about using the following Conda command to reproduce a computational environment?

`conda env export -f environment.yaml`

A)  Only explicitly installed packages will be exactly replicated on another computer with a different OS   
B)  Your exported environment can be exactly reproduced on another computer with any operating system   
C)  This command allows Conda to automatically resolve OS-related package conflicts   
D)  Your exported environment can be exactly reproduced on another computer with the same operating system    
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q16 = ...
```

```python
grader.check("q16")
```

**Step 10:** Export your environment into `environment.yaml` file in way that will allow your environment to work **on different operating systems**. Make sure it is saved into your `project repo`

<!-- BEGIN QUESTION -->


<div class="alert alert-info">
    
### Question 17

Suppose that we would like to provide an environment file called `environment.yaml` to our coworkers so that they can easily replicate our computational environment on their machines for a particular project.

However, you want to **manually** modify some parameters in your `environment.yaml` to make sure that:

- The environment name is `project-env` 
- It installs Python 3.10 or higher 
- It installs the latest version of libraries that you need to run the `Report.ipynb` that you can find in your `project repo`. Note that `sklearn` is a short of `scikit-learn` package
- The environment will be available to JupyterLab 
    
Save your changes push and commit them to the remote repo.  Create a pull request to merge the content of `developer-1` branch into the `main` branch, and merge the changes.

Then navigate to the remote `project` repo, open the `environment.yaml` and provide a link to the file.

The link should look something like this: https://github.com/YOUR-USER-NAME/project/blob/main/environment.yaml


</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Step 11:** Let's make final touches that we learned in previous Modules to make sure that your `project` is ready for submission:
    
- Modify your `README` file so it reflect what is you repo is about
- Add necessary files to `.gitignore`
    
    
Push all the changes to the remote `project` repo


<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 18

Send us a URL link to your `project` repo (the link should look similar to this one: https://github.com/YOUR_USER_NAME/project)

You will be evaluated based on:

- project structure
- file/folder naming
- presence of all required file
- comprehensive README.md
- .gitignore file
- quality of commit message

</div>

_Points:_ 4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-danger" style="color:black">
    
**Restart and run all before submitting**
    
Before submitting,
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see the output of all the cells properly.
You can do this by clicking the ▶▶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for this course,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

------

Assignment 6 completed, Congratulations! 🤪

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 1)
```
>>> assert 'a' in q1.lower()
>>> assert 'c' in q1.lower()
>>> assert 'f' in q1.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'b' not in q1.lower()
>>> assert 'd' not in q1.lower()
>>> assert 'e' not in q1.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q10 (points: 1)
```
>>> assert 'b' in q10.lower()
>>> assert 'c' in q10.lower()
>>> assert 'd' in q10.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q10.lower()
>>> assert 'e' not in q10.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q11 (points: 1)
```
>>> assert 'a' == q11.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q12 (points: 1)
```
>>> assert 'b' == q12.lower()
```
failure_message: You are missing one of the correct answers


Tests for q13 (points: 1)
```
>>> assert 'b' in q13.lower()
>>> assert 'd' in q13.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q13.lower()
>>> assert 'c' not in q13.lower()
>>> assert 'e' not in q13.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q14 (points: 1)
```
>>> q14_string = q14.replace(" ", "")
>>> q14_string = q14_string.replace('"pkgs/main"',"'pkgs/main'")
>>> assert q14_string == "condasearchscikit-learn|grep'pkgs/main'|tail-n1"
```
failure_message: Your answer is incorrect


Tests for q15 (points: 1)
```
>>> q15_string = q15.replace(" ", "")
>>> assert q15_string == "condainstallscikit-learn=1.2.1-cpkgs/main"
```
failure_message: Your answer is incorrect


Tests for q16 (points: 1)
```
>>> assert 'd' == q16.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q2 (points: 1)
```
>>> assert 'd' in q2.lower()
>>> assert 'b' in q2.lower()
>>> assert 'c' in q2.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q2.lower()
```
failure_message: At least one of the options you have selected is not correct
```
>>> assert q2.upper() == 'DBC'
```
failure_message: At least one of the options you have selected is not in correct order


Tests for q3 (points: 1)
```
>>> assert 'c' == q3.lower()
```
failure_message: You chose incorrect asnwer


Tests for q4 (points: 1)
```
>>> assert 'd' == q4.lower()
>>>
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q5 (points: 1)
```
>>> assert 'b' == q5.lower()
```
failure_message: You are missing one of the correct answers


Tests for q6 (points: 1)
```
>>> assert 'd' == q6.lower()
```
failure_message: You are missing one of the correct answers


Tests for q7 (points: 1)
```
>>> assert q7.upper() == 'BACBAC'
```
failure_message: At least one of the options you have selected is not in correct order


Tests for q8 (points: 1)
```
>>> assert 'd' == q8.lower()
```
failure_message: You chose incorrect asnwer


Tests for q9 (points: 1)
```
>>> assert 'a' == q9.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers
